// Numbas version: exam_results_page_options
{"name": "Complex number in polar form", "extensions": [], "custom_part_types": [], "resources": [["question-resources/polar2.png", "/srv/numbas/media/question-resources/polar2.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Complex number in polar form", "tags": [], "metadata": {"description": "<p></p>\n<p></p>", "licence": "None specified"}, "statement": "<p>Complex Numbers</p>\n<p></p>\n<p></p>", "advice": "<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">To convert $z=-\\frac{1}{3}\\left(\\sqrt{3}+1\\right)+\\frac{1}{3}\\left(\\sqrt{3}-1\\right)i$ into polar form, we need to find the modulus $r$, and the argument $\\theta$. Then the polar form is written as</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">\\[z=a+ib=r(\\cos{(\\theta)}+i\\sin{(\\theta)}) \\]</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\"><img src=\"resources/question-resources/polar2.png\" width=\"436\" height=\"278\"/></p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">The modulus is given by</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">\\begin{align} r&amp;=\\sqrt{\\left(\\frac{\\sqrt{3}+1}{3}\\right)^2+\\left(\\frac{\\sqrt{3}-1}{3}\\right)^2}\\\\ &amp;= \\frac{1}{3}\\sqrt{8}\\\\ &amp;=\\frac{2\\sqrt{2}}{3}.\\end{align}</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">For the argument, we have</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">\\begin{align} \\tan{\\theta}&amp;= \\frac{(\\sqrt{3}-1)/3}{-(\\sqrt{3}+1)/3} \\\\ &amp;=\\frac{(1-\\sqrt{3})(1-\\sqrt{3})}{(1+\\sqrt{3})(1-\\sqrt{3})}\\\\ &amp;=\\sqrt{3}-2 .\\end{align}</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">Since the real part of $z$ is negative, and the imaginary part of $z$ is posative, \\begin{align} \\theta &amp;= \\pi+\\tan^{-1}{(\\sqrt{3}-2)} \\\\ &amp;= \\pi-\\frac{\\pi}{12} \\\\ &amp;= \\frac{11\\pi}{12} .\\end{align}</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">Therefore, the polar form is</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">\\[ -\\frac{1}{3}\\left(\\sqrt{3}+1\\right)+\\frac{1}{3}\\left(\\sqrt{3}-1\\right)i=\\frac{2\\sqrt{2}}{3}\\left(\\cos{\\left(\\frac{11\\pi}{12}\\right)}+i\\sin{\\left(\\frac{11\\pi}{12}\\right)}\\right).\\]</p>\n<p></p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Express the following complex number in polar form</p>\n<p>\\[ -\\frac{1}{3}\\left(\\sqrt{3}+1\\right)+\\frac{1}{3}\\left(\\sqrt{3}-1\\right)i\\]</p>", "answer": "2sqrt(2)/3(cos(11pi/12)+i sin(11pi/12))", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}]}], "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}