// Numbas version: exam_results_page_options
{"name": "Complex number in Cartesian form", "extensions": [], "custom_part_types": [], "resources": [["question-resources/polar2.png", "/srv/numbas/media/question-resources/polar2.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Complex number in Cartesian form", "tags": [], "metadata": {"description": "<p></p>\n<p></p>", "licence": "None specified"}, "statement": "<p>Complex Numbers</p>\n<p></p>\n<p></p>", "advice": "<p>Using $i^2=-1$, we have</p>\n<p>\\begin{align}\\frac{-5\\sqrt{3}+5i}{(-i+1)^2-2\\sqrt{3}} &amp;= \\frac{-5\\sqrt{3}+5i}{(1-2i+i^2)-2\\sqrt{3}}\\\\ &amp;=\\frac{(-5\\sqrt{3}+5i)(-2\\sqrt{3}+2i)}{(-2\\sqrt{3}-2i)(-2\\sqrt{3}+2i)}\\\\ &amp;=\\frac{20}{16}-i\\frac{20\\sqrt{3}}{16}\\\\&amp;=\\frac{5}{4}-i\\frac{5\\sqrt{3}}{4}.\\end{align}</p>\n<p></p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Express the following complex number in cartesian form $a+ib$</p>\n<p>\\[ \\frac{-5\\sqrt{3}+5i}{(-1+i)^2-2\\sqrt{3}} \\]</p>", "answer": "5/4-(5 sqrt(3))/(4)i", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}]}], "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}