// Numbas version: exam_results_page_options
{"name": "Complex roots", "extensions": [], "custom_part_types": [], "resources": [["question-resources/polar2.png", "/srv/numbas/media/question-resources/polar2.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Complex roots", "tags": [], "metadata": {"description": "<p></p>\n<p></p>", "licence": "None specified"}, "statement": "<p>Complex Numbers</p>\n<p></p>\n<p></p>", "advice": "<p>Let $z^4=w$,&nbsp; then&nbsp; $3z^4+\\bar{z}^4=4$ becomes $3w+\\bar{w}=4$. Now, let $w=a+ib$ so $\\bar{w}=a-ib$. Then</p>\n<p>\\begin{align}3w+\\bar{w}&amp;= 3(a+ib)+(a-ib)\\\\ &amp;= 4a+2bi\\end{align}</p>\n<p>Thus $4a+2bi=4$, giving $a=1$ and $b=0$, so $w=a+ib=1$. Hence $z^4=1$, so solving for $z$ gives $z=i, -i,-1, 1.$</p>\n<p></p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Find all complex roots of the following equation</p>\n<p>\\[ 3z^4+\\bar{z}^4=4 \\]</p>\n<p>where $\\bar{z}$ denotes the complex conjugate of $z$.</p>\n<p>Write your answer in square brackets, separating the solutions with a comma (e.g. if your solutions are $z=1$ and $z=2$, write [1,2]).</p>", "alternatives": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "alternativeFeedbackMessage": "", "useAlternativeFeedback": false, "answer": "[-i, i, -1, 1]", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "alternativeFeedbackMessage": "", "useAlternativeFeedback": false, "answer": "[1, -1, i, -i]", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "alternativeFeedbackMessage": "", "useAlternativeFeedback": false, "answer": "[-1, 1, -i, i]", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "3", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "alternativeFeedbackMessage": "", "useAlternativeFeedback": false, "answer": "[-i, -1, 1, i]", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "answer": "[i, -i, -1, 1]", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}]}], "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}