// Numbas version: exam_results_page_options
{"name": "T2 W7 Q2 1920", "extensions": [], "custom_part_types": [], "resources": [["question-resources/polar2.png", "/srv/numbas/media/question-resources/polar2.png"], ["question-resources/Screen_Shot_2019-02-13_at_01.51.45.png", "/srv/numbas/media/question-resources/Screen_Shot_2019-02-13_at_01.51.45.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "T2 W7 Q2 1920", "tags": [], "metadata": {"description": "<p></p>\n<p></p>", "licence": "None specified"}, "statement": "<p>Mathematical Methods (T2 Week 7)</p>", "advice": "<p>Solutions to Question 1:</p>\n<p>Part(a)</p>\n<p>(i) Let $I:=\\int x\\sqrt{1+x}dx$. Use the substitution $u=1+x$, so that $x=u-1$, $dx=du$. So \\begin{align}I &amp;=\\int (u-1)u^{1/2}du=\\int(u^{3/2}-u^{1/2})du\\nonumber\\\\ &amp; = \\frac{2}{5}u^{5/2}-\\frac{2}{3}u^{3/2}+C\\nonumber\\\\ &amp;=\\frac{2}{15}(1+x)^{3/2}(3x-2)+C\\end{align}</p>\n<p>(Note: we use \":=\" to say that the right hand side is being named/defined as the left hand side. Now that we've defined it, from now on, we just use $I=...$.)</p>\n<p>(ii) Let $I:=\\int \\frac{x}{(x^2+2)^3}dx$. Use the substitution $u=x^2$, so that $du=2xdx,$</p>\n<p>So \\begin{align} I &amp;= \\frac{1}{2}\\int\\frac{1}{(u+2)^3}du=\\frac{1}{2}\\int (u+2)^{-3}du\\nonumber \\\\ &amp;= \\frac{1}{4}(u+2)^{-2}+C\\nonumber\\\\ &amp; =-\\frac{1}{4(x^2+2)^2}+C \\end{align}</p>\n<p></p>\n<p>(iii) Let $I := \\int \\frac{1}{\\cos^2\\left(\\frac{1}{4} x-\\frac{1}{4}\\right)}dx = \\int \\sec^2{\\left(\\frac{1}{4} x-\\frac{1}{4}\\right)}dx$</p>\n<p>Using $u=\\frac{1}{4}x-\\frac{1}{4}$, so that $dx=4du$, we have $I=4\\int \\sec^2(u)du = 4\\tan(u)+C.$ So \\[I = 4\\tan\\left(\\frac{x-1}{4}\\right)+C.\\]</p>\n<p></p>\n<p>(iv) Let $I := \\int x^4\\cos(x^5+3)dx$</p>\n<p>Using $u=x^5,$ so $du=5x^4dx$, we have</p>\n<p>\\[I=\\frac{1}{5}\\int \\cos(u+3)du = \\frac{1}{5}\\sin(u+3)+C.\\]</p>\n<p>So $I = \\frac{1}{5}\\sin(x^5+3)+C.$</p>\n<p></p>\n<p>(v) Let $I: = \\int\\frac{\\sin(\\sqrt{x})}{\\sqrt{x}}dx$.</p>\n<p>Using $u=\\sqrt{x}$, so that $du = \\frac{1}{2}x^{-1/2}dx$, we have $I=2\\int \\sin(u)du = -2\\cos(u)+C = -2\\cos(\\sqrt{x})+C$.</p>\n<p></p>\n<p>(vi) Let $I:=\\int \\frac{1}{x\\sqrt{\\ln(x)}}dx$. Using $u=\\ln(x)$, so that $du=\\frac{dx}{x}$, we have</p>\n<p>$I = \\int\\frac{1}{\\sqrt{u}}du=2u^{1/2}+C=2\\sqrt{\\ln(x)}+C$</p>\n<p></p>\n<p>(vii) Let $I:=\\int \\frac{e^x}{1+e^x}dx$. Using $u=e^x$, so $du=e^xdx$, we have $I=\\int\\frac{du}{1+u}=\\ln(1+u)+C=\\ln(1+e^x)+C.$</p>\n<p></p>\n<p>(viii) Let $I:=\\int \\frac{\\ln(x)}{x} dx$. Using $u=\\ln(x)$, so that $du = \\frac{1}{x}dx$, we have $I = \\int udu=\\frac{1}{2}u^2+C=\\frac{1}{2}(\\ln(x))^2+C$.</p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Evaluate the following integrals.</p>\n<p>(Please use an upper case C to indicate a constant.)</p>\n<p>(i) $\\int x\\sqrt{1+x}dx$ = [[0]]</p>\n<p></p>\n<p>(ii) $\\int \\frac{x}{(x^2+2)^3}dx$ = [[1]]</p>\n<p></p>\n<p>(iii) $\\int \\frac{1}{\\cos^2\\left(\\frac{1}{4}x-\\frac{1}{4}\\right)} dx$ = [[2]]</p>\n<p></p>\n<p>(iv) $\\int x^4\\cos(x^5+3)dx$ = [[3]]</p>\n<p></p>\n<p>(v) $\\int \\frac{\\sin(\\sqrt{x})}{\\sqrt{x}}dx$ = [[4]]</p>\n<p></p>\n<p>(vi) $\\int\\frac{1}{x\\sqrt{\\ln (x)}}dx$ = [[5]]</p>\n<p></p>\n<p>(vii) $\\int \\frac{e^x}{1+e^x}dx$ = [[6]]</p>\n<p></p>\n<p>(viii) $\\int \\frac{\\ln (x)}{x}dx$ = [[7]]</p>", "gaps": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "1", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "2/15(1+x)^(3/2)(3x-2)+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "-(1)/(4(x^2+2)^2)+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "4*tan((x-1)/4)+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "1/5*sin(x^5+3)+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "-2*cos(sqrt(x))+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "2*sqrt(ln(x))+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "ln(1+e^x)+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}, {"type": "jme", "useCustomName": false, "customName": "", "marks": "1.5", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "answer": "1/2*(ln(x))^2+C", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": [{"name": "c", "value": ""}, {"name": "x", "value": ""}]}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}]}], "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}