// Numbas version: exam_results_page_options
{"name": "Complex number in Cartesian form (2)", "extensions": [], "custom_part_types": [], "resources": [["question-resources/polar2.png", "/srv/numbas/media/question-resources/polar2.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Complex number in Cartesian form (2)", "tags": [], "metadata": {"description": "<p></p>\n<p></p>", "licence": "None specified"}, "statement": "<p>Complex Numbers</p>\n<p></p>\n<p></p>", "advice": "<p>Using $i^2=-1$, we have</p>\n<p>\\begin{align} \\frac{(1+2i)(2+i)}{(1+3i)^2-2i(2+i)}&amp; = \\frac{2-2+5i}{1+6i-9-4i+2}\\\\ &amp;= \\frac{5i}{-6+2i}\\\\ &amp;=\\frac{5i(-6-2i)}{(-6+2i)(-6-2i)}\\\\ &amp;=\\frac{10-30i}{40}\\\\ &amp;=\\frac{1}{4}-\\frac{3}{4}i \\end{align}</p>\n<p></p>", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "jme", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">Express the following complex number in Cartesian form $a+ib$</p>\n<p style=\"background-color: transparent; color: #000000; font-family: Verdana,Arial,Helvetica,sans-serif; font-size: 14px; font-style: normal; font-variant: normal; font-weight: 400; letter-spacing: normal; min-height: 0px; orphans: 2; text-align: left; text-decoration: none; text-indent: 0px; text-transform: none; -webkit-text-stroke-width: 0px; white-space: normal; word-spacing: 0px;\">\\[ \\frac{(1+2i)(2+i)}{(1+3i)^2-2i(2+i)} \\]</p>", "answer": "1/4-3/4*i", "showPreview": true, "checkingType": "absdiff", "checkingAccuracy": 0.001, "failureRate": 1, "vsetRangePoints": 5, "vsetRange": [0, 1], "checkVariableNames": false, "singleLetterVariables": false, "allowUnknownFunctions": true, "implicitFunctionComposition": false, "caseSensitive": false, "valuegenerators": []}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}]}], "contributors": [{"name": "Layal Hakim", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2345/"}]}